# Exercício interativo

Antes da submissão final do seu problema, certifique que tudo funciona como esperado. Primeiro, **reinicie o kernel** (Kernel$\rightarrow$Restart) e depois **execute todas as células** (Cell$\rightarrow$Run All).

Preencha todos lugares que digam `SUA SOLUÇÃO` or `YOUR ANSWER HERE`, e o seu nome e dos colaboradores (outros alunos no grupo) abaixo.

```python
NAME = ""
COLLABORATORS = ""
```

---

Este teste utiliza o Jupyter notebook.

![](jupyter.png)

## Instruções

* A prova deve ser resolida individualmente
* Provas com cola/plágio serão anuladas
* Você pode consultar a internet ou a documentação do Python, mas não pode copiar as respostas.
* Será realizada uma etapa posterior para detecção de plágio. Provas que contenham plágio serão anuladas.

---
## Lista aninada (2,0 pts)

Crie uma função que recebe uma lista que pode conter valores ou outras listas dentro. A função deve remover todo aninhamento e retornar uma lista plana de valores.

<!-- grade_id: flat, points: 0, grade: false, solution: true -->
```python
def flat(lst):
    """Retorna uma lista sem aninhamento"""
    NotImplemented # escreva sua solução aqui...
```

Célula de testes

```python
flat([1, 2, [3, 4]])
```

<!-- grade_id: flat-1, points: 1, grade: true, solution: false -->
```python
"""Remove um grau de aninhamento."""
assert flat([1, 2, 3]) == [1, 2, 3]
assert flat([1, 2, [3, 4]]) == [1, 2, 3, 4]
assert flat([1, 2, [3, 4], 5]) == [1, 2, 3, 4, 5]
print('Parabéns!')
```

<!-- grade_id: flat-2, points: 1, grade: true, solution: false -->
```python
"""Remove n graus de aninhamento."""
assert flat([1, 2, 3]) == [1, 2, 3]
assert flat([1, 2, [3, 4]]) == [1, 2, 3, 4]
assert flat([1, 2, [3, 4], 5]) == [1, 2, 3, 4, 5]
assert flat([1, [2, [3, [4]], 5]]) == [1, 2, 3, 4, 5]
print('Parabéns!')
```

---
## Romanos vs arábicos (2,0 pts)

Crie um programa que converta números romanos para números arábicos

| Símbolo | Valor             |
|---------|-------------------|
| I       | 1 (unus)          |
| V       | 5 (quinque)       |
| X       | 10 (decem)        |
| L       | 50 (quinquaginta) |
| C       | 100 (centum)      |
| D       | 500 (quingenti)   |
| M       | 1,000 (mille)     |


<!-- grade_id: roman, points: 0, grade: false, solution: true -->
```python
NUMEROS = {'I': 1, 'V': 5, 'X': 10, 'L': 50, 'C': 100, 'D': 500, 'M': 1000}

def arabic(roman):
    """Converte número romano para arábico"""
    NotImplemented # escreva sua solução aqui...
```

Teste o ano atual arabic('MMXIX') == 2019

```python
arabic('MMXIX')
```

<!-- grade_id: roman-1, points: 1, grade: true, solution: false -->
```python
"""Verifica valores simples"""
assert arabic('C') == 100
assert arabic('XIII') == 13
assert arabic('MDCLXVI') == 1666
assert arabic('MMDCCLXXVII') == 2777
print('Parabéns!')
```

<!-- grade_id: roman-2, points: 1, grade: true, solution: false -->
```python
"""Verifica casos complexos"""
assert arabic('IV') == 4
assert arabic('XLII') == 42
assert arabic('XCII') == 92
assert arabic('CDXLII') == 442
assert arabic('CMXCII') == 992

print('Parabéns!')
```

---
## Maior nota (1,0pts)

A função melhor_nota() recebe uma lista de pares com (nome, nota) e retorna o nome associado à maior nota.

<!-- grade_id: maior-nota, points: 0, grade: false, solution: true -->
```python
def maior_nota(lst):
    """Retorna o nome associado à maior nota"""
    NotImplemented # escreva sua solução aqui...
```

Teste sua implementação aqui. 

```python
maior_nota([('Isaac', 9.5), ('Charles', 9.2), ('Olavo', 0.1), ('Albert', 9.8),])
```

<!-- grade_id: maior-nota-testes, points: 1, grade: true, solution: false -->
```python
"Testa a maior nota"

assert maior_nota([('Isaac', 9.5), ('Charles', 9.2), ('Olavo', 0.1), ('Albert', 9.8),]) == 'Albert'
assert maior_nota([('Foo', 0)]) == 'Foo'
print('Parabéns!')
```

---
## Que tipo de X é você? (5,0 pts)

Vamos criar um programa que simula um quiz de internet. 

Primeiramente, um pouco de matemática. Ainda que vários quizzes simplesmente retornam resultados aleatórios, é possível fornecer as respostas com um pouco mais de precisão. Suponha que existam 3 perguntas com três alternatiavs cada e 3 possíveis resultados. Podemos definir um sistema de pontuação onde cada resultado pode receber uma pontuação para cada possível resposta. 

<!-- grade_id: perfis, points: 0, grade: true, solution: true -->
```python
NotImplemented # escreva sua solução aqui...
```

O primeiro passo para resolver nosso problema é criar uma função que recebe uma lista de votos e um perfil (representado como uma lista de dicionários) e calcula a pontuação para aquele perfil.

<!-- grade_id: quizz-pontos, points: 0, grade: false, solution: true -->
```python
def calcula_pontos(votos, perfil):
    NotImplemented # escreva sua solução aqui...
```

```python
# Verifique sua implementação
calcula_pontos([1, 0, 1, 2, 2], [{0: 1}, {0: 1}, {2: 1}, {0: 1}, {1: 1}])
```

<!-- grade_id: quizz-pontos-teste, points: 1.5, grade: true, solution: false -->
```python
"Verifica implementação de calcula_pontos"
assert calcula_pontos([1, 0, 1, 2, 2], [{0: 1}, {0: 1}, {2: 1}, {0: 1}, {1: 1}]) == 1
assert calcula_pontos([0, 0, 2, 0, 1], [{0: 1}, {0: 1}, {2: 1}, {0: 1}, {1: 1}]) == 5
assert calcula_pontos([0, 1, 2, 0, 1], [{0: 1}, {0: 1, 1: -0.5}, {2: 1}, {0: 1}, {1: 1}]) == 3.5
print('Parabéns!')
```

Agora vamos criar uma função que percorre todos os perfis no dicionário PERFIS e retorne qual é o perfil mais semelhante ao padrão de votos dado. DICA: use a função calcula pontos para calcular o ponto de cada perfil e retorne o nome do perfil com a maior pontuação.

<!-- grade_id: quizz-perfil, points: 0, grade: false, solution: true -->
```python
def obtem_perfil(votos):
    NotImplemented # escreva sua solução aqui...
```

```python
# Verifique sua implementação
obtem_perfil([1, 2, 0])
```

<!-- grade_id: quizz-perfil-test, points: 1.5, grade: true, solution: false -->
```python
"Verifica implementação de obtem_perfil()"
assert obtem_perfil([2, 1, 0]) == PERFIS[0][0]
assert obtem_perfil([1, 0, 1]) ==  PERFIS[1][0]
assert obtem_perfil([0, 2, 2]) ==  PERFIS[2][0]

print('Parabéns!')
```

Complete a solução. Crie um programa que realiza as 3 perguntas e calcula o perfil do usuário. Você não deve alterar a conteúdo dos dicionários na variável PERFIS, mas você pode (e deve) alterar os nomes de cada perfil. Desta forma, escolha as respostas de acordo com os valores já salvos nesta variável.

Você deve escolher algum tema qualquer (que tipo de personagem do Game of Thrones, que tipo de político Brasileiro, que tipo de personagem do Star Wars, etc) e desenvolver o programa a partir daí.

<!-- grade_id: quizz-perfil-full, points: 2, grade: true, solution: true -->
```python
# Função auxiliar para perguntar alternativas
def pergunta(texto, alternativas):
    print(texto)
    print('  a)', alternativas[0])
    print('  b)', alternativas[1])
    print('  c)', alternativas[2])
    escolha = input('Escolha: ').lower()
    if escolha not in 'abc':
        print('Opção inválida. Tente novamente.\n')
        return pergunta(texto, alternativas)
    return 'abc'.index(escolha)


# Começe seu programa modificando 
print('Olá! Vamos perguntar algumas coisas para saber que tipo de X você é')

perfil = 'desconhecido'
opt_1 = pergunta('Como você faria Y?', ['De um jeito', 'De outro jeito', 'De um terceiro jeito'])
# Continue a sua solução aqui...
NotImplemented # escreva sua solução aqui...

print('\nParabéns! Vocé é um %s!' % perfil)
```