# Exercício interativo

Antes da submissão final do seu problema, certifique que tudo funciona como esperado. Primeiro, **reinicie o kernel** (Kernel$\rightarrow$Restart) e depois **execute todas as células** (Cell$\rightarrow$Run All).

Preencha todos lugares que digam `SUA SOLUÇÃO` or `YOUR ANSWER HERE`, e o seu nome e dos colaboradores (outros alunos no grupo) abaixo.

```python
NAME = ""
COLLABORATORS = ""
```

---

Este teste utiliza o Jupyter notebook.

![](jupyter.png)

## Instruções

* Teste pode ser resolvido individualmente ou em dupla
* Não pode colar/plagiar outras duplas
* Você pode consultar a internet ou a documentação do Python, mas não pode copiar as respostas.
* Será realizada uma etapa posterior para detecção de plágio. Provas que contenham plágio serão anuladas.

---
## Média e variância (2,0 pts)

Crie uma função que receba uma lista de valores e retorne uma dupla com a média e o desvio padrão do conjunto.
O desvio padrão é a raiz quadrada da variância, que pode ser calculada como a soma dos quadrados menos o quadrado da soma dividido pelo número de elementos.

<!-- grade_id: stats, points: 0, grade: false, solution: true -->
```python
from math import sqrt


def stats(valores):
    """Retorna a (media, desvpad) do conjunto de números fornecido."""
    NotImplemented # escreva sua solução aqui...
    return media, desvpad
```

A média e desvio padrão dos números [1, 2, 3] são respectivamente 2 e 0.8165. Verifique o resultado.

```python
stats([1, 2, 3])
```

<!-- grade_id: stats-test, points: 2, grade: true, solution: false -->
```python
"""Verifica se a função foi implementada corretamente."""
assert stats([1, 1, 1]) == (1, 0)
assert stats([-1, 0, 1]) == (0, sqrt(2/3))
assert stats(range(100)) == (49.5, 28.86607004772212)
print('Parabéns!')
```

---
## Sistema de registro (4,0)

Um sistema possui um módulo de autenticação em que os usuários podem iniciar uma seção após a verificação de senha e e-mail. As informações dos usuários são armazenadas na lista USUARIOS. Você deve implementar as funções de login() e recuperar_senha() para passar nos testes abaixo.

<!-- grade_id: registro, points: 0, grade: false, solution: true -->
```python
# Não edite a lista de usuários
USUARIOS = [
    {
        'nome': 'Ada Lovelace',
        'emails': ['ada@gmail.co.uk', 'algo-ada@hotmail.com'],
        'nascimento': (1815, 12, 10),
        'senha': 'algoritmo',
    },
    {
        'nome': 'Alan Turing',
        'emails': ['turing@gmail.co.uk'],
        'nascimento': (1912, 6, 23),
        'senha': 'enigma',
    },
    {
        'nome': 'Grace Hopper',
        'emails': ['amazing-grace@usnavy.gov'],
        'nascimento': (1906, 12, 9),
        'senha': 'cobol',
    },
    {
        'nome': 'Donald Knuth',
        'emails': ['knuth@gmail.com', 'art-of-programming@caltech.edu'],
        'nascimento': (1938, 1, 10),
        'senha': 'LaTeX',
    },
]

def login(email, senha):
    """
    Retorna um usuário caso o e-mail e senha estejam corretos. Retorna None caso contrário.
    """
    NotImplemented # escreva sua solução aqui...
    return None


def recuperar_senha(email, nascimento):
    """
    Mostra a senha do usuário ou uma mensagem de erro caso o e-mail seja 
    inválido ou a data de nascimento não corresponda à correta.
    """
    NotImplemented # escreva sua solução aqui...
    return 'Data de nascimento ou e-mail inválidos'
```

<!-- grade_id: registro-login-1, points: 1, grade: true, solution: false -->
```python
"Testes para a função de login()"

assert login('turing@gmail.co.uk', 'enigma') == USUARIOS[1]
assert login('turing@gmail.co.uk', 'password-incorreto') is None
assert login('ada@gmail.co.uk', 'algoritmo') == USUARIOS[0]
print('Parabéns!')
```

<!-- grade_id: registro-login-2, points: 1, grade: true, solution: false -->
```python
"Testes para a função de login() - múltiplos e-mails"

assert login('knuth@gmail.com', 'LaTeX') == USUARIOS[-1]
assert login('art-of-programming@caltech.edu', 'LaTeX') == USUARIOS[-1]
assert login('art-of-programming@caltech.edu', 'latex') is None
print('Parabéns!')
```

<!-- grade_id: registro-recuperar-1, points: 1, grade: true, solution: false -->
```python
"Testes para a função recuperar_senha() - casos de sucesso"

assert recuperar_senha('turing@gmail.co.uk', (1912, 6, 23)) == 'A senha é "enigma"'
assert recuperar_senha('amazing-grace@usnavy.gov', (1906, 12, 9)) == 'A senha é "cobol"'
print('Parabéns!')
```

<!-- grade_id: registro-recuperar-2, points: 1, grade: true, solution: false -->
```python
"Testes para a função recuperar_senha() - casos de erro"

assert recuperar_senha('turing@gmail.co.uk', (1901, 2, 30)) == 'Data de nascimento ou e-mail inválidos'
assert recuperar_senha('bad@email.com.br', (1912, 6, 23)) == 'Data de nascimento ou e-mail inválidos'
print('Parabéns!')
```

---
## Índice Gini (4,0)

O índice [Gini](http://www.ipea.gov.br/desafios/index.php?option=com_content&id=2048:catid=28) é uma medida de desigualdade de renda ou outros recursos criada pelo matemático italiano Conrado Gini. Seu valor varia de 0 (distribuição completamente equalitária) até 1 (uma única pessoa concentra todos recursos). Vamos criar uma função que calcula o índice Gini a partir de uma lista de valores.

Como esta tarefa é um pouco complicada, este exercício foi dividido em várias partes.

![Curva Gini](https://s1.static.brasilescola.uol.com.br/img/2015/06/grafico-do-indice-de-gini.jpg)

### Índice Gini Riqueza acumulada

O primeiro passo para o cálculo do índice Gini é obter a lista de riquezas acumuladas. Para isso, crie uma função chamada `soma_acc()` que recebe uma lista de valores e recebe o valor das somas acumuladas. Considere a lista de entrada `[a, b, c, ...]`. A saída das somas acumuladas é dada por `[a, a + b, a + b + c, ...]`.

<!-- grade_id: gini-acc, points: 0, grade: false, solution: true -->
```python
def soma_acc(valores):
    """
    Retorna as somas parciais da lista de valores
    """
    NotImplemented # escreva sua solução aqui...
```

<!-- grade_id: gini-acc-tests, points: 1.5, grade: true, solution: false -->
```python
"Testes para a função recuperar_senha() - casos de erro"

assert soma_acc([1, 2, 3, 4]) == [1, 3, 6, 10]
assert soma_acc([4, 3, 2, 1]) == [4, 7, 9, 10]
assert soma_acc([42]) == [42]
print('Parabéns!')
```

### Gini: área embaixo da curva

Normalizando tanto o eixo x quanto o eixo y de 0 até 1, a área embaixo da reta de total igualdade é igual a 0.5. Podemos calcular a área da curva da riqueza acumulada (curva de Lorenz) aproximando a integral por uma soma usando o ponto final de cada intervalo como a altura de cada elemento na curva

![Aproximação da integral](https://upload.wikimedia.org/wikipedia/commons/4/45/RightRiemann2.svg)

Obs.: lembre-se que o primeiro elemento começa em (0, 0) e diz que zero porcento da riqueza foi acumulado pelos
zero porcento mais pobre da população.

Crie uma função que recebe uma lista de valores e retorne a área abaixo da curva de Lorenz.

<!-- grade_id: gini-area, points: 0, grade: false, solution: true -->
```python
def area_curva(valores):
    """
    Calcula a área da curva de Lorenz para o conjunto de 
    valores dados.
    """
    NotImplemented # escreva sua solução aqui...
```

<!-- grade_id: gini-area-testes, points: 1.5, grade: true, solution: false -->
```python
"Testa a área abaixo da curva de Lorenz"

assert area_curva([1, 1, 1]) == area_curva([2, 2, 2, 2]) == 0.5
assert area_curva([1, 2, 3, 4]) == 0.4
assert area_curva([0, 0, 0, 0, 0, 0, 1]) == 0.125
print('Parabéns!')
```

### Gini: cálculo do coeficiente

Para finalizar temos que fazer apenas alguns ajustes. Primeiramente, a lista de valores de entrada deve ser
ordenada (utilize a função `sorted` do Python). Em seguida, retornamos o coeficiente Gini como a área entre as duas curvas dividida pela área abaixo da reta de total igualdade.

<!-- grade_id: cell-3162f7f61b141ea9, points: 0, grade: false, solution: true -->
```python
def gini(valores):
    """
    Calcula o coeficiente Gini a partir dos valores.
    """
    NotImplemented # escreva sua solução aqui...
```

<!-- grade_id: gini-especiais, points: 0.5, grade: true, solution: false -->
```python
"Gini, casos especiais"
assert gini([1, 1, 1, 1]) == 0
assert gini([0, 0, 0, 0, 0, 0, 1]) == 0.75
assert gini([1, 0, 0, 0, 0, 0, 0]) == 0.75
print('Parabéns!')
```

<!-- grade_id: gini-exemplos, points: 0.5, grade: true, solution: false -->
```python
"Gini, exemplos"

assert abs(gini([1, 2, 3, 4]) - 0.20) < 0.1
assert abs(gini([1, 1, 1, 2, 4, 8, 16, 32, 64]) - 0.59) < 0.1  # ~Brasil
print('Parabéns!')
```